## Remember that the solutions need to be saved int variable provided.

```python
# Import the needed libraries
import numpy as np
print("hey")
new = np.arange(8)
print(new)
```

Output:
```
hey
[0 1 2 3 4 5 6 7]
```

**13. Create a 10x10 array with random values and find the minimum and maximum values**

<!-- grade_id: cell-d37d68df72162149, points: 0, grade: false, solution: true -->
```python
# 1 point
sol1 = np.random.normal(size=(10,10))
min = np.amin(sol1)
max = np.max(sol1)
print(min,max)
#raise NotImplementedError()
```

Output:
```
-1.7916045388284758 1.9726873966471332
```

<!-- grade_id: cell-f1d4fa08074f4dcd, points: 1, grade: true, solution: false -->
```python
assert sol1.shape == (10,10)
```

**15. Create a 2d 10x10 array with 1 on the border and 0 inside**

<!-- grade_id: cell-6584e6c544771f0c, points: 0, grade: false, solution: true -->
```python
#1 points


new_arr = np.zeros( (10,10) )
new_arr[0] = 1
new_arr[9] = 1
new_arr[:,0] = 1
new_arr[:,9] = 1
sol2 = new_arr
#raise NotImplementedError()
```

<!-- grade_id: cell-9429aab0f4477a86, points: 1, grade: true, solution: false -->
```python
assert sol2.shape == (10,10)
assert sol2[0].all() == 1
assert sol2[:,0].all() == 1
assert sol2[1][1] == 0
assert sol2[2][2] == 0
assert sol2[3][3] == 0
assert sol2[4][4] == 0
assert sol2[1][2] == 0
assert sol2[2][3] == 0
assert sol2[3][4] == 0
assert sol2[4][5] == 0
```

**34. How to get all the dates corresponding to the month of July 2016**

<!-- grade_id: cell-9784b014be5fdf5a, points: 0, grade: false, solution: true -->
```python
#2 points
sol3 = 0

# YOUR CODE HERE
sol3 = np.arange('2016-07','2016-08',dtype='datetime64[D]')
print(sol3)
```

Output:
```
['2016-07-01' '2016-07-02' '2016-07-03' '2016-07-04' '2016-07-05'
 '2016-07-06' '2016-07-07' '2016-07-08' '2016-07-09' '2016-07-10'
 '2016-07-11' '2016-07-12' '2016-07-13' '2016-07-14' '2016-07-15'
 '2016-07-16' '2016-07-17' '2016-07-18' '2016-07-19' '2016-07-20'
 '2016-07-21' '2016-07-22' '2016-07-23' '2016-07-24' '2016-07-25'
 '2016-07-26' '2016-07-27' '2016-07-28' '2016-07-29' '2016-07-30'
 '2016-07-31']
```

<!-- grade_id: cell-8beeae591e9b50ed, points: 2, grade: true, solution: false -->
```python
assert sol3.shape == (31,)
assert type(sol3[0]) is np.datetime64
```

**42. Consider two random arrays based on the normal and uniform distributions 10x10 sol41 and sol42, check if they are equal**

<!-- grade_id: cell-5000744434ae05a2, points: 0, grade: false, solution: true -->
```python
#2 points
sol41 = np.random.uniform(size=(10,10))
sol42 = np.random.uniform(size=(10,10))
comparisson = sol41 == sol42
print(comparisson)
```

Output:
```
[[False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]
 [False False False False False False False False False False]]
```

<!-- grade_id: cell-807bf6c3ff3f3b0d, points: 2, grade: true, solution: false -->
```python
bol = True

for l in comparisson:
    for c in l:
        if c:
            bol = False


assert bol
```

**72. How to swap two rows of an array?** 

<!-- grade_id: cell-079026f268b5211c, points: 0, grade: false, solution: true -->
```python
#3 points
sol5 = np.array([[1,2,3],[4,5,6]])

(sol5[[0,1]]) = sol5[[1,0]]
# YOUR CODE HERE

```

<!-- grade_id: cell-e4f878f0864f33bd, points: 3, grade: true, solution: false -->
```python
assert sol5[0][0] == 4
assert sol5[0][1] == 5
assert sol5[0][2] == 6

assert sol5[1][0] == 1
assert sol5[1][1] == 2
assert sol5[1][2] == 3

```

**75. How to compute averages using a sliding window over an array?**

<!-- grade_id: cell-ccafbda4e4cb7996, points: 0, grade: false, solution: true -->
```python
arr = np.array([1,2,3,4,5,6,7,8,9,10])

# YOUR CODE HERE
sol6 = np.cumsum(arr)/len(arr)



```

<!-- grade_id: cell-1983fc0078a860ff, points: 3, grade: true, solution: false -->
```python
assert sol6[-1] == 5.5
assert sol6[0] == 0.1
assert sol6[4] == 1.5
```

```python

```