## Imports

```python
# pip install --upgrade scikit-learn
```

```python
import time
from IPython.display import clear_output
import numpy    as np
import pandas   as pd
import seaborn  as sb
import matplotlib.pyplot as plt
import sklearn  as skl

from sklearn import pipeline      # Pipeline
from sklearn import preprocessing # OrdinalEncoder, LabelEncoder
from sklearn import impute
from sklearn import compose
from sklearn import model_selection # train_test_split
from sklearn import metrics         # accuracy_score, balanced_accuracy_score, plot_confusion_matrix
from sklearn import set_config

set_config(display='diagram') # Useful for display the pipeline

print("Pandas  ", pd.__version__)
print("Sklearn ", skl.__version__) # Try to use 0.24
```

Output:
```
Pandas   1.3.3
Sklearn  0.24.2
```

## Get the dataset
- **CLOUD = True**: Download dataset from Kaggle. Necesary for cloud enviroments like COLAB. **Specify your [kaggle credentials](https://www.kaggle.com/docs/api)**.
- **CLOUD = False**: Get the dataset from your local machine. **Specify the data path**.

```python
DATA_PATH = "../../Datasets/Tabular/titanic/"
```

## Load data

```python
df      = pd.read_csv("train.csv", index_col='PassengerId')
df_test = pd.read_csv("test.csv",  index_col='PassengerId')

print("Train DataFrame:", df.shape)
print("Test DataFrame: ", df_test.shape)
```

Output:
```
Train DataFrame: (891, 11)
Test DataFrame:  (418, 10)
```

## Check missings

```python
df.isnull().sum()
```

Output:
```
Survived      0
Pclass        0
Name          0
Sex           0
Age         177
SibSp         0
Parch         0
Ticket        0
Fare          0
Cabin       687
Embarked      2
dtype: int64
```

```python
df_test.isnull().sum()
df_test.Name
```

Output:
```
PassengerId
892                                 Kelly, Mr. James
893                 Wilkes, Mrs. James (Ellen Needs)
894                        Myles, Mr. Thomas Francis
895                                 Wirz, Mr. Albert
896     Hirvonen, Mrs. Alexander (Helga E Lindqvist)
                            ...                     
1305                              Spector, Mr. Woolf
1306                    Oliva y Ocana, Dona. Fermina
1307                    Saether, Mr. Simon Sivertsen
1308                             Ware, Mr. Frederick
1309                        Peter, Master. Michael J
Name: Name, Length: 418, dtype: object
```

# Exercise 1 (2pts):
Extract the title (Mr, Mrs, ... ) from the "Name" column.

Tips:
- split(',')[1] to get the 2nd part, and remove the surnamename
- split('.')[0] to get the 1str part, and remove the name

<!-- grade_id: cell-d56107842f0f54d2, points: 0, grade: false, solution: true -->
```python
# CODE HERE get_Title_from_Name funtion
# Create this function using lambda (not def)

get_Title_from_Name = lambda x : x.split(',')[1].split('.')[0].strip()

### BEGIN SOLUTION
df['Title'] = df['Name'].apply(get_Title_from_Name)
### END SOLUTION
df_test['Title'] = df_test['Name'].apply(get_Title_from_Name)

```

<!-- grade_id: cell-f3d989a17018344b, points: 2, grade: true, solution: false -->
```python
assert df['Title'].values[0] == "Mr"
assert df['Title'].values[1] == "Mrs"
assert df['Title'].values[2] == "Miss"

assert df_test['Title'].values[0] == "Mr"
assert df_test['Title'].values[1] == "Mrs"
assert df_test['Title'].values[414] == "Dona"
```

# Exercise 2 (1pts):
Apply the title_dictionary to get a better information about the title. You have to overwrite the Title variable.

```python
title_dictionary = {
    "Capt": "Officer",
    "Col": "Officer",
    "Major": "Officer",
    "Jonkheer": "Royalty",
    "Don": "Royalty",
    "Sir" : "Royalty",
    "Dr": "Officer",
    "Rev": "Officer",
    "the Countess":"Royalty",
    "Mme": "Mrs",
    "Mlle": "Miss",
    "Ms": "Mrs",
    "Mr" : "Mr",
    "Mrs" : "Mrs",
    "Miss" : "Miss",
    "Master" : "Master",
    "Lady" : "Royalty"
}
```

```python
df['Title'].values[886]
```

Output:
```
'Rev'
```

<!-- grade_id: cell-5c6a842c812fafda, points: 0, grade: false, solution: true -->
```python
# Use map to apply the prevous dict
df['Title'] = df['Title'].map(title_dictionary)
df_test["Title"] = df_test["Title"].map(title_dictionary)

### BEGIN SOLUTION

### END SOLUTION

```

```python
df.Ticket.value_counts()
```

Output:
```
347082      7
CA. 2343    7
1601        7
3101295     6
CA 2144     6
           ..
9234        1
19988       1
2693        1
PC 17612    1
370376      1
Name: Ticket, Length: 681, dtype: int64
```

<!-- grade_id: cell-669f5a637e57e835, points: 1, grade: true, solution: false -->
```python
assert df['Title'].values[886] == "Officer"
assert df_test['Title'].values[417] == "Master"
```

# Exercise OPTINAL (0pts):
Try to extract some information from the feature **Ticket**. Search on Internet if that colum has some kind of information.

```python

```

# Exercise OPTIONAL (0pts):
Try to extract some information from the feature **Cabin**. Search on Internet if that colum has some kind of information.

```python
df
```

Output:
```
Survived  Pclass  \
PassengerId                     
1                   0       3   
2                   1       1   
3                   1       3   
4                   1       1   
5                   0       3   
...               ...     ...   
887                 0       2   
888                 1       1   
889                 0       3   
890                 1       1   
891                 0       3   

                                                          Name     Sex   Age  \
PassengerId                                                                    
1                                      Braund, Mr. Owen Harris    male  22.0   
2            Cumings, Mrs. John Bradley (Florence Briggs Th...  female  38.0   
3                                       Heikkinen, Miss. Laina  female  26.0   
4                 Futrelle, Mrs. Jacques Heath (Lily May Peel)  female  35.0   
5                                     Allen, Mr. William Henry    male  35.0   
...                    
```

# Preprocessing
For X data, notice that...
- We drop Survived because is the target variable
- We drop Name because we have extracted the Title: Mr, Mrs, ...
- We drop Ticket because it has no information -> see df.Ticket.nunique()
- We drop Cabin because it has a lot of missings (77% are missings)

Then, we identify **numerical** variables and **categorical** variables,

```python
x = df.drop(columns=["Survived", 'Name', 'Ticket', 'Cabin']) # X DATA (WILL BE TRAIN+VALID DATA)
y = df["Survived"] # 0 = No, 1 = Yes

x_test = df_test.drop(columns=['Name', 'Ticket', 'Cabin']) # # X_TEST DATA (NEW DATA)
```

```python
cat_vars  = ['Sex', 'Embarked', 'Title']         # x.select_dtypes(include=[object]).columns.values.tolist()
num_vars  = ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age'] # x.select_dtypes(exclude=[object]).columns.values.tolist()

print("\nNumerical features:\n", num_vars)
print("\nCategorical features:\n", cat_vars)
```

Output:
```
Numerical features:
 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']

Categorical features:
 ['Sex', 'Embarked', 'Title']
```

# Exercise 3 (2pts):
Create a **ColumnTransformer for Tree Models**. You need to create 2 pipelines (one for numerical and other for categories). Remember:
- Categorical pipeline: Some SimpleImputer -> Some Encoder
- Numerical pipeline: Some SimpleImputer -> NO Encoder

<!-- grade_id: cell-c607e75fb38ea248, points: 0, grade: false, solution: true -->
```python
from sklearn.impute import SimpleImputer
from sklearn.pipeline import Pipeline,make_pipeline
from sklearn.preprocessing import OneHotEncoder
from sklearn.compose import ColumnTransformer
"""
num_preprocessing = pipeline.Pipeline(steps=[
  # Some SimpleImputer here
])

cat_preporcessing = pipeline.Pipeline(steps=[
  # Some SimpleImputer here
  # Some Encoder here. Remember to handle_unknown
])

tree_prepro = compose.ColumnTransformer(transformers=[
    ('num', num_preprocessing, num_vars),
    ('cat', cat_preporcessing, cat_vars),
], remainder='drop') # Drop other vars not specified in num_vars or cat_vars

tree_prepro
""";
num_imputer = SimpleImputer(strategy='most_frequent')
cat_imputer = SimpleImputer(strategy='most_frequent')
encoder = OneHotEncoder(handle_unknown='ignore')
num_preprocessing = Pipeline([('imputer',num_imputer)])
cat_preprocessing = Pipeline([('imputer',cat_imputer),('encoder',encoder)])
tree_prepro = ColumnTransformer([('num',num_preprocessing,num_vars),('cat',cat_preprocessing,cat_vars)],remainder='drop')
### BEGIN SOLUTION
# Drop other vars not specified in num_vars or cat_vars



### END SOLUTION

```

<!-- grade_id: cell-e636558135fb5975, points: 2, grade: true, solution: false -->
```python
assert type(tree_prepro)      is compose._column_transformer.ColumnTransformer
assert type(num_preprocessing) is pipeline.Pipeline
assert type(cat_preprocessing) is pipeline.Pipeline
assert len(num_preprocessing) == 1
assert len(cat_preprocessing) == 2
```

# Exercise 4 (1pts):
1. Complete the diccionary with some Tree Models.
2. Then we put each model in a Pipeline where:
   - first is the prepocessing with the column Transformer
   - Then is the Tree model
3. Display the fullpipeline of the LGBMClassifier

```python
from sklearn.tree          import DecisionTreeClassifier
from sklearn.ensemble      import RandomForestClassifier
from sklearn.ensemble      import ExtraTreesClassifier
from sklearn.ensemble      import AdaBoostClassifier
from sklearn.ensemble      import GradientBoostingClassifier
from sklearn.experimental  import enable_hist_gradient_boosting # Necesary for HistGradientBoostingClassifier
from sklearn.ensemble      import HistGradientBoostingClassifier
from xgboost               import XGBClassifier
from lightgbm              import LGBMClassifier
from catboost              import CatBoostClassifier
```

<!-- grade_id: cell-76a3be8223730c5a, points: 0, grade: false, solution: true -->
```python
"""
tree_classifiers = {
  "Decision Tree": DecisionTreeClassifier(),
  "Extra Trees": ExtraTreesClassifier()
  "Random Forest":RandomForestClassifier()
  "AdaBoost":
  "Skl GBM":
  "Skl HistGBM":
  "XGBoost":
  "LightGBM":
  "CatBoost":
tree_classifiers = {name: pipeline.make_pipeline(tree_prepro, model) for name, model in tree_classifiers.items()}
tree_classifiers["LightGBM"]
"""
classifiers = {
"Decision Tree":DecisionTreeClassifier(),
"Extra Trees":RandomForestClassifier(),
"Random Forest":ExtraTreesClassifier(),
"AdaBoost":AdaBoostClassifier(),
"Skl GBM":GradientBoostingClassifier(),
"Skl HistGBM":HistGradientBoostingClassifier(),
"XGBoost":XGBClassifier(),
"LightGBM":LGBMClassifier(),
"CatBoost":CatBoostClassifier(verbose=False)
}



tree_classifiers = {name:make_pipeline(tree_prepro,model) for name,model in classifiers.items()}
from sklearn import set_config
set_config(display='diagram')
tree_classifiers['AdaBoost']


```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(strategy='most_frequent'))]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(strategy='most_frequent')),
                                                                  ('encoder',
                                                                   OneHotEncoder(handle_unknown='ignore'))]),
                                                  ['Sex', 'Embarked',
                                                   'Title'])])),
  
```

<!-- grade_id: cell-d4744022b37e2e9b, points: 1, grade: true, solution: false -->
```python
for pipe in tree_classifiers.values():
    assert type(pipe) is pipeline.Pipeline
```

# Exercise 5 (3pts):
Define a simple split validation strategy with:
- 80% for train
- 20% for validation
- With stratification
- random_state=0

And train all the models in a for loop

<!-- grade_id: cell-64f17e0d448bca7e, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
from time import time
"""
x_train, x_val, y_train, y_val = model_selection.train_test_split(
    # CODE HERE
)
"""
### BEGIN SOLUTION
results=[]
x_train,x_val,y_train,y_val = train_test_split(x,y,train_size=0.8,stratify=y)

for model_name,model in tree_classifiers.items():
    start = time()
    model.fit(x_train,y_train)
    pred = model.predict(x_val)
    
    total_time = time()-start
    results.append({
        "Model": model_name,
        "Accuracy": metrics.accuracy_score(y_val,pred)*100,
        "Bal Acc." : metrics.balanced_accuracy_score(y_val,pred)*100,
        "Time":total_time,
    })
results=pd.DataFrame(results)



results
```

Output:
```
/Users/saidalikhonalikhonov/opt/anaconda3/envs/new_env/lib/python3.9/site-packages/xgboost/sklearn.py:888: UserWarning: The use of label encoder in XGBClassifier is deprecated and will be removed in a future release. To remove this warning, do the following: 1) Pass option use_label_encoder=False when constructing XGBClassifier object; and 2) Encode your labels (y) as integers starting with 0, i.e. 0, 1, 2, ..., [num_class - 1].
  warnings.warn(label_encoder_deprecation_msg, UserWarning)
[14:30:41] WARNING: /opt/concourse/worker/volumes/live/7a2b9f41-3287-451b-6691-43e9a6c0910f/volume/xgboost-split_1619728204606/work/src/learner.cc:1061: Starting in XGBoost 1.3.0, the default evaluation metric used with the objective 'binary:logistic' was changed from 'error' to 'logloss'. Explicitly set eval_metric if you'd like to restore the old behavior.
           Model   Accuracy   Bal Acc.      Time
0  Decision Tree  79.329609  77.779974  0.017989
1    Extra Trees  82.122905  80.322793  0.113312
```

<!-- grade_id: cell-af27b33fe88ad384, points: 3, grade: true, solution: false -->
```python
assert results["Accuracy"].min() > 75
assert results["Bal Acc."].min() > 75
assert len(results) == 9
```

# Exercise 6 (3pts):
Define a 10 Fold cross validation strategy with:
- With stratification
- shuffle=True
- random_state=0

And train all the models in a for loop.

Tip you can use **[cross_val_predict](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.cross_val_predict.html)** for both training and predict with 

<!-- grade_id: cell-37bf16dc25b43732, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import StratifiedKFold,cross_validate
"""
skf = model_selection.StratifiedKFold(
    # CODE HERE
)
"""
### BEGIN SOLUTION
results_ord=[]
skf = StratifiedKFold(n_splits=10)
### END SOLUTION
for model_name,model in tree_classifiers.items():
    start=time()
    cross_validate(model,x,y,cv=skf)
    pred = model.predict(x_val)
    total_time= time()-start
    results_ord.append({
        "Model": model_name,
        "Accuracy": metrics.accuracy_score(y_val,pred)*100,
        "Bal Acc." : metrics.balanced_accuracy_score(y_val,pred)*100,
        "Time":total_time,
    })
    
"""
for model_name, model in tree_classifiers.items():
    start_time = time.time()
        
    # TRAIN AND GET PREDICTIONS USING cross_val_predict() and x,y
    pred = # CODE HERE

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y_val, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y_val, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                              
                              
"""

### BEGIN SOLUTION

### END SOLUTION

```

Output:
```
/Users/saidalikhonalikhonov/opt/anaconda3/envs/new_env/lib/python3.9/site-packages/xgboost/sklearn.py:888: UserWarning: The use of label encoder in XGBClassifier is deprecated and will be removed in a future release. To remove this warning, do the following: 1) Pass option use_label_encoder=False when constructing XGBClassifier object; and 2) Encode your labels (y) as integers starting with 0, i.e. 0, 1, 2, ..., [num_class - 1].
  warnings.warn(label_encoder_deprecation_msg, UserWarning)
/Users/saidalikhonalikhonov/opt/anaconda3/envs/new_env/lib/python3.9/site-packages/xgboost/sklearn.py:888: UserWarning: The use of label encoder in XGBClassifier is deprecated and will be removed in a future release. To remove this warning, do the following: 1) Pass option use_label_encoder=False when constructing XGBClassifier object; and 2) Encode your labels (y) as integers starting with 0, i.e. 0, 1, 2, ..., [num_class - 1].
  warnings.warn(label_encoder_deprecation_msg, UserWarning)
/Users/saidali
```

```python
results_ord=pd.DataFrame(results_ord)
best=results_ord.sort_values(by='Accuracy',ascending=False).iloc[0,0]


```

Output:
```
str
```

<!-- grade_id: cell-57cc085faad513cf, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 75
assert len(results_ord) == 9
```

# Exercise 7.1
Train with all data the best model

```python
best=results_ord.sort_values(by='Accuracy',ascending=False).iloc[0,0]
best_pipe = tree_classifiers[best]

best_pipe.fit(x,y)

# Fit best model with all data

### BEGIN SOLUTION

### END SOLUTION
```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(strategy='most_frequent'))]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(strategy='most_frequent')),
                                                                  ('encoder',
                                                                   OneHotEncoder(handle_unknown='ignore'))]),
                                                  ['Sex', 'Embarked',
                                                   'Title'])])),
  
```

# Exercise 7.2 (2pts)
With your best model, generate the predicitions for test data (x_test)

<!-- grade_id: cell-44b93a0dbd4eb6fc, points: 0, grade: false, solution: true -->
```python
# test_pred = # Get the predictions for x_test
test_pred = best_pipe.predict(x_test)
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-a29fc8d24284520e, points: 2, grade: true, solution: false -->
```python
assert len(test_pred) == 418
assert np.unique(test_pred).tolist() == [0,1]
```

# Exercise 7.3

Submit to kaggle.

- You can use the kaggle command line app. Check https://github.com/Kaggle/kaggle-api

```python
sub = pd.DataFrame(test_pred, index=x_test.index, columns=["Survived"])
sub
```

Output:
```
Survived
PassengerId          
892                 0
893                 1
894                 0
895                 0
896                 1
...               ...
1305                0
1306                1
1307                0
1308                0
1309                1

[418 rows x 1 columns]
```

```python
sub.to_csv("sub.csv")
```

```python
!kaggle competitions submit -c titanic -f sub.csv -m "My submission message"
```

Output:
```
100%|████████████████████████████████████████| 2.77k/2.77k [00:02<00:00, 955B/s]
Successfully submitted to Titanic - Machine Learning from Disaster
```